# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
answer = None
# Assess CMS files and create column for the years 2014,2015,2016
cms2014 = pd.read_csv('/data/cms2014.csv')
cms2014.columns
cms2014['year']=2014
cms2015 = pd.read_csv('/data/cms2015.csv')
cms2015['year']=2015
cms2016 = pd.read_csv('/data/cms2016.csv')
cms2016['year'] = 2016


### Solution
```

```python
print(cms2014.shape)
print(cms2015.shape)
print(cms2016.shape)
```

```python
# Make sure columns match between years 2014 and 2015 and then years 2015 and 2016
cms2014.columns==cms2015.columns
```

```python
# Make sure columns match between years 2015 and 2016 
cms2016.columns==cms2015.columns
```

```python
# 2016 columns
cms2016.columns
```

```python
#Column names are the same but some are lower case and some are capita; change everything to lower so everything is the same
cms2016.columns = cms2016.columns.str.strip().str.lower()
cms2015.columns = cms2015.columns.str.strip().str.lower()
cms2014.columns = cms2014.columns.str.strip().str.lower()
```

```python
cms2015.columns == cms2016.columns
```

```python
# Add all the columns together by append 
cms=cms2014.append(cms2015).append(cms2016)
```

```python
cms.shape
```

```python
# Create the answer object and assign the cms dataframe to it
answer = cms
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
# Assert Checks
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total#_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **/data/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
# Read the csv file 
answer = None
npi= pd.read_csv('/data/npidata.csv', float_precision='round_trip')
npi.head()

### SOLUTION


```

```python
cms.shape 
```

```python
# Should create a year for 2014,2015,2016
cms.year.unique()
```

```python
#Create Matching key field
npi['npi'] = npi. NPI
```

```python
pd.options.display.float_format = '{:.10f}'.format
```

```python
# Left join using merge so that all cms records and npi records that match are shown
answer = cms.merge(npi, on = ['npi'], how = 'left')
```

```python
# answer shape
answer.shape 
```

```python
# Coserve Memory
del npi
del cms
import gc 
```

```python
gc.collect()
```

```python
import pandas as pd 
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
# Creating a new column 'State' and creating a column called 'XX' if it doesn't have any information 


answer.loc[answer['Provider Business Mailing Address State Name'].isnull() == True, 'State'] = 'XX'
answer.loc[answer['Provider Business Mailing Address State Name'].isnull() == False, 'State'] = answer['Provider Business Mailing Address State Name']

### SOLUTION


```

```python
# Describe total services
answer['total_services'].describe()
```

```python
answer=pd.pivot_table(answer, values= 'total_services', index='State', columns='year', aggfunc= 'sum')
```

```python
print(answer)
```

```python
pd.options.display.float_format = '{:.7f}'.format
answer.sum()
```

```python
answer.head()
```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
# Display output inline by executing code

%matplotlib inline


### SOLUTION

```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

```python
%matplotlib inline
answer.plot.bar()
```

## Submit your work as usual

```python

```

```python

```