# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
from pathlib import Path
HOME = str(Path.home())

ALLERGIES_FILE="/data/allergies.json"
```

```python
systems = json.load(open(ALLERGIES_FILE))
```

```python
systems
```

Output:
```
{'resourceType': 'Bundle',
 'type': 'searchset',
 'total': 3,
 'link': [{'relation': 'self',
   'url': 'https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance?patient=Tbt3KuCY0B5PSrJvCu2j-PlK.aiHsu2xUjUM8bWpetXoB'}],
 'entry': [{'fullUrl': 'https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TBwnNbrAqC0Qw5Ha7AFT-2AB',
   'link': [{'relation': 'self',
     'url': 'https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TBwnNbrAqC0Qw5Ha7AFT-2AB'}],
   'search': {'mode': 'match'},
   'resource': {'resourceType': 'AllergyIntolerance',
    'recordedDate': '2015-08-24T23:11:36Z',
    'status': 'confirmed',
    'criticality': 'CRITL',
    'id': 'TBwnNbrAqC0Qw5Ha7AFT-2AB',
    'onset': '2012-11-07T00:00:00Z',
    'recorder': {'display': 'MOORE, NICK',
     'reference': 'https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Practitioner/TItWfhjChtlo0pFh9nzctSQB'},
    'patient': {'display': 'Jason Argonaut',
     'reference': 'https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Pat
```

```python
type(systems)
```

Output:
```
dict
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
import json
### BEGIN SOLUTION
'''This json file will open the file, read the json and return a number of entries in the file'''
def allergy_count(json_file):
    f = open(HOME + json_file)
    data = json.load(f)
    return len(data["entry"])

    
### END SOLUTION
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
### BEGIN SOLUTION
'''This json file will count how many unique patients that are in this structure'''
def patients_count(json_file):
    patients = []
    data = open(HOME + json_file)
    for entry in data["entry"]:
        patient = entry("resource", "patient", "display")
        if patient not in patients:
                    patients.extend(patient)
    return len(patients)
    print(patient)
    print(lens(patients))

### END SOLUTION
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-9-0809ae29512d> in <module>
----> 1 patient_count(ALLERGIES_FILE)

NameError: name 'patient_count' is not defined
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
### BEGIN SOLUTION
'''This json file will document what each allergy a patient has.  A patient may not have any allergies or a patient may have more than one allergy.'''
def allergy_per_patient(json_file):
        allergies = []
        f= open(HOME + json_file)
        data = json.load(f)
        for entry in data["entry"]:
            patient = entry["resource"]["patient"]["display"]
        if patient not in allergies:
            print("patient has no allergies")
        elif  allergies[patient] > 2:
            print("patient has more than 2 allergies")
        else:
            allergies[patient] += 1
            print("patient has 1 or 2 allergies")
            return allergies
            print(allergies)
            print(lens(patient))

### END SOLUTION
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
'''This json file will output patients name along with allergy and type of reaction'''
def allergy_list(json_file):
    
    allergies = []
    
    f = open(HOME + json_file)
    data = json.load(f)
    
    for entry in data["entry"]:
        allergies.append([entry["resource"]["patient"]["display"], entry["resource"]["substance"]["text"], entry["resource"]["reaction"][0]["manifestation"][0]["text"]])
   
    return allergies


### END SOLUTION
                         

```

```python
allergies = [
{
  "resourceType": "Bundle",
  "type": "searchset",
  "total": 3,
  "link": [
    {
      "relation": "self",
      "url": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance?patient=Tbt3KuCY0B5PSrJvCu2j-PlK.aiHsu2xUjUM8bWpetXoB"
    }
  ],
  "entry": [
      
    {
      "fullUrl": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TBwnNbrAqC0Qw5Ha7AFT-2AB",
      "link": [
        {
          "relation": "self",
          "url": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TBwnNbrAqC0Qw5Ha7AFT-2AB"
        }
      ],
      "search": {
        "mode": "match"
      },
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-08-24T23:11:36Z",
        "status": "confirmed",
        "criticality": "CRITL",
        "id": "TBwnNbrAqC0Qw5Ha7AFT-2AB",
        "onset": "2012-11-07T00:00:00Z",
        "recorder": {
          "display": "MOORE, NICK",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Practitioner/TItWfhjChtlo0pFh9nzctSQB"
        },
        "patient": {
          "display": "Jason Argonaut",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Patient/Tbt3KuCY0B5PSrJvCu2j-PlK.aiHsu2xUjUM8bWpetXoB"
        },
        "substance": {
          "text": "PENICILLIN G",
          "coding": [
            {
              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",
              "code": "7980",
              "display": "PENICILLIN G"
            },
            {
              "system": "http://fdasis.nlm.nih.gov",
              "code": "Q42T66VG0C",
              "display": "PENICILLIN G"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2012-11-07T00:00:00Z",
            "manifestation": [
              {
                "text": "Hives"
              }
            ],
            "note": {
              "text": "Severity low enough to be prescribed if needed."
            }
          }
        ],
        "note": {
          "text": "Severity low enough to be prescribed if needed."
        }
      }
    },

          {
      "fullUrl": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TBwnNbrAqC0Qw5Ha3AFT-2AB",
      "link": [
        {
          "relation": "self",
          "url": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TBwnNbrAqC0Qw5Ha3AFT-2AB"
        }
      ],
      "search": {
        "mode": "match"
      },
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2011-09-14T23:11:36Z",
        "status": "confirmed",
        "criticality": "CRITL",
        "id": "TBwnNbrAqC0Qw5Ha3AFT-2AB",
        "onset": "2012-11-07T00:00:00Z",
        "recorder": {
          "display": "JONES, ALICE",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Practitioner/TItWfhjChtlo0pFh9nzctSQA"
        },
        "patient": {
          "display": "Paul Boal",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Patient/Tbt3KuCY0B5PSrJvCu2j-PlK.aiHsu2xUjUM8bWpetXoZ"
        },
        "substance": {
          "text": "PENICILLIN G",
          "coding": [
            {
              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",
              "code": "7980",
              "display": "PENICILLIN G"
            },
            {
              "system": "http://fdasis.nlm.nih.gov",
              "code": "Q42T66VG0C",
              "display": "PENICILLIN G"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2009-11-07T00:00:00Z",
            "manifestation": [
              {
                "text": "Bruising"
              },
              {
                "text": "Shortness of breath"
              }
            ],
            "note": {
              "text": "Severe. Do not prescribe."
            }
          }
        ],
        "note": {
          "text": "Severe. Do not prescribe."
        }
      }
    },
    {
      "fullUrl": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TPcWiBG2h2E114Vh0sRT8fQB",
      "link": [
        {
          "relation": "self",
          "url": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TPcWiBG2h2E114Vh0sRT8fQB"
        }
      ],
      "search": {
        "mode": "match"
      },
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-11-07T20:55:10Z",
        "status": "confirmed",
        "criticality": "CRITL",
        "id": "TPcWiBG2h2E114Vh0sRT8fQB",
        "onset": "2010-05-02T00:00:00Z",
        "recorder": {
          "display": "MOORE, NICK",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Practitioner/TItWfhjChtlo0pFh9nzctSQB"
        },
        "patient": {
          "display": "Jason Argonaut",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Patient/Tbt3KuCY0B5PSrJvCu2j-PlK.aiHsu2xUjUM8bWpetXoB"
        },
        "substance": {
          "text": "SHELLFISH-DERIVED PRODUCTS",
          "coding": [
            {
              "system": "http://hl7.org/fhir/ndfrt",
              "code": "N0000007624",
              "display": "SHELLFISH-DERIVED PRODUCTS"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2010-05-02T00:00:00Z",
            "manifestation": [
              {
                "text": "Itching"
              }
            ]
          }
        ]
      }
    },
      
    {
      "fullUrl": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TKebKfLXzu6Sp.LY-IpvpmQB",
      "link": [
        {
          "relation": "self",
          "url": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/AllergyIntolerance/TKebKfLXzu6Sp.LY-IpvpmQB"
        }
      ],
      "search": {
        "mode": "match"
      },
      "resource": {
        "resourceType": "AllergyIntolerance",
        "recordedDate": "2015-11-07T20:56:34Z",
        "status": "confirmed",
        "criticality": "CRITH",
        "id": "TKebKfLXzu6Sp.LY-IpvpmQB",
        "onset": "2014-03-07T00:00:00Z",
        "recorder": {
          "display": "MOORE, NICK",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Practitioner/TItWfhjChtlo0pFh9nzctSQB"
        },
        "patient": {
          "display": "Jason Argonaut",
          "reference": "https://open-ic.epic.com/FHIR/api/FHIR/DSTU2/Patient/Tbt3KuCY0B5PSrJvCu2j-PlK.aiHsu2xUjUM8bWpetXoB"
        },
        "substance": {
          "text": "STRAWBERRY",
          "coding": [
            {
              "system": "http://www.nlm.nih.gov/research/umls/rxnorm",
              "code": "892484",
              "display": "STRAWBERRY"
            },
            {
              "system": "http://fdasis.nlm.nih.gov",
              "code": "4J2TY8Y81V",
              "display": "STRAWBERRY"
            }
          ]
        },
        "reaction": [
          {
            "certainty": "confirmed",
            "onset": "2014-03-07T00:00:00Z",
            "manifestation": [
              {
                "text": "Anaphylaxis"
              }
            ]
          }
        ]
      }
    }
  ]
}
]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list(ALLERGIES_FILE) == output

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json
'''For this json file, I will be taking three parameters in consideration for the type of allergic reaction a patient may have to what type of substance they are allergic too.'''
### BEGIN SOLUTION
def allergy_reaction(json_file, patient, substance):
    allergies = allergy_list(json_file)
    for name, sub, react in allergy_list(json_file):
            if name == patient and sub == substance:
                return react
    return None
### END SOLUTION
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---
---

# Stretch (Extra) Problem

Work on either of the stretch problems below can earn you up to 25 free points toward the midterm assignment.  That is, if you complete one of these extra problems successfully, you can skip 1 of the problems that will appear on the midterm exam coming up next week.

The midterm will be distributed 10/14 and due 10/24.



---
---

### STRETCH for October 2022 - For those looking for an additional challenge

As I've mentioned in class, CMS is now enforcing a rule around price transparency.  Every facility that take Medicare payments is required to publish a "machine readable" file with it's pricing infomration for a number of common procedures across all of the payers they work with.  There are two examples of such files in the `/data/` directory: `whiteriver.json` and `saline.xml`.

If you want to compare contracted prices across these two hospitals, you'll need to read in the information from both of those files into some kind of data structure, then merge the data together from those two files.  See what you can do.

See if you can create an output file that has the following fields:
* HOSPITAL
* PROCEDURE_CODE
* PAYER
* AMOUNT

If you choose to work on this, you may get stuck at some point and you won't know if you're _doing it right_. Make some assumptions. Document your questions in this notebook.



```
Procedure Code |  Description  |  Gross Charges  |  Aetna  |  QualChoice
```

```python

```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.